## Question 1: Basic SQL (13 pts)
In this problem, you'll interact with a toy SQL database using Python's
built-in `sqlite3` package. Documentation can be found at
<https://docs.python.org/3/library/sqlite3.html>. For this problem,
we'll use a popular toy SQLite database, called Chinook, which
represents a digital music collection. See the documentation at <https://github.com/lerocha/chinook-database/blob/master/README.md>
for a more detailed explanation. We'll use the `chinook.sqlite` file:

```python
import sqlite3
con = sqlite3.connect('chinook.sqlite')
cur = con.cursor()
```

**1(a)** (2 pt) Load the database using the Python `sqlite3` package. How many tables are in the database? Save the answer in the variable `n_tables`.

<!-- grade_id: 1a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
tables = cur.execute("SELECT count(*) FROM sqlite_master WHERE type='table'")
n_tables = tables.fetchone()[0]
```

<!-- grade_id: load_ds, points: 2, grade: true, solution: false -->
```python
assert n_tables == 11
```

**1(b)** (2 pts) What are the names of the tables in the database? Save the answer as
    a set of strings, `table_names`. **Note:** you should write Python `sqlite3`
    code to answer this; don't just look up the answer in the
    documentation!

<!-- grade_id: 1b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
table_names = set()
for table in cur.execute("SELECT name FROM sqlite_master WHERE type='table'"):
    table_names.add(table[0])
print(table_names)
```

Output:
```
{'Track', 'Invoice', 'PlaylistTrack', 'Artist', 'Album', 'Playlist', 'Customer', 'MediaType', 'InvoiceLine', 'Employee', 'Genre'}
```

<!-- grade_id: correct_1b, points: 2, grade: true, solution: false -->
```python
expected = {'Album', 'Genre', 'Playlist', 'PlaylistTrack', 'Employee', 'Customer', 'InvoiceLine', 'Track', 'Artist', 'MediaType', 'Invoice'}
assert table_names == expected
```

**1(c)** (3 pts) Write a function `albums_starting_with(c)` that takes as an argument a single character `c` and
    returns a list of the primary keys (AlbumIds) of all the albums whose titles
    start with that character. Your function should ignore case, so that
    the inputs "a" and "A" yield the same results. Include error
    checking that raises an error in the event that the input is not a
    single character.


<!-- grade_id: grade_13, points: 0, grade: false, solution: true -->
```python
def albums_starting_with(c):
    # YOUR CODE HERE
    if not isinstance(c, str) or len(c) != 1:
        raise ValueError("input must be a single character")
    query = "SELECT AlbumId FROM Album WHERE Title LIKE ?"
    cur.execute(query, (f"{c}%",))
    results = [row[0] for row in cur.fetchall()]
    return results
```

<!-- grade_id: correct_13, points: 3, grade: true, solution: false -->
```python
res = albums_starting_with('a')
assert type(res) == list
assert len(res) == 32
for i in [10, 14, 15, 24, 26, 29, 74, 167, 319]:
    assert i in res

try:
    albums_starting_with(3)
except Exception:
    pass
else:
    raise Exception("should raise an exception")
    
```

**1(d)** (3 pts) Write a function `songs_starting_with(c)` that takes as an argument a single character and
    returns a list of the primary keys (TrackIds) of all the **songs** whose album
    names begin with that letter. Again, your function should ignore
    case and perform error checking as in the previous exercise (again ignoring case).
    **Hint:** you'll need a JOIN statement here. Don't forget that you
    can use the `cursor.description` attribute to find out about tables
    and the names of their columns.

<!-- grade_id: grade_4, points: 0, grade: false, solution: true -->
```python
def songs_starting_with(c):
    # YOUR CODE HERE
    if not isinstance(c, str) or len(c) != 1:
        raise ValueError("input must be a single character")
    query = """
        SELECT t.TrackId 
        FROM Track t
        JOIN Album a ON t.AlbumId = a.AlbumId
        WHERE a.Title LIKE ?
        """
    cur.execute(query, (f"{c}%",))
    results = [row[0] for row in cur.fetchall()]
    return results
```

<!-- grade_id: load_4, points: 3, grade: true, solution: false -->
```python
res = songs_starting_with('a')
assert type(res) == list
assert len(res) == 369
for i in [85, 86, 87, 331, 332, 333, 923, 924, 925]:
    assert i in res

try:
    songs_starting_with(3)
except Exception:
    pass
else:
    raise Exception("should raise an exception")

try:
    songs_starting_with('res')
except Exception:
    pass
else:
    raise Exception("should raise an exception")
```

**1(e)** (3 pts) Write a function `cost_of(c)` that takes as an argument a single character and
    returns the cost of buying every song (consider only the songs that were sold - you need to look into InvoiceLine table) whose album begins with that
    letter. This cost should be based on the tracks' unit prices when it was sold, so
    that the cost of buying a set of tracks is simply the sum of the
    unit prices of all the tracks in the set. Again your function should
    ignore case and perform appropriate error checking.


<!-- grade_id: grade_5, points: 0, grade: false, solution: true -->
```python
def cost_of(c):
    # YOUR CODE HERE
    if not isinstance(c, str) or len(c) != 1:
        raise ValueError("input must be a single character")
    query = """
        SELECT SUM(track_price)
        FROM (
            SELECT il.TrackId, MIN(il.UnitPrice) AS track_price
            FROM InvoiceLine AS il
            INNER JOIN Track  AS t ON t.TrackId  = il.TrackId
            INNER JOIN Album  AS a ON a.AlbumId  = t.AlbumId
            WHERE UPPER(a.Title) LIKE UPPER(? || '%')
            GROUP BY il.TrackId
        ) AS per_track
    """
    cur.execute(query, (c,))
    result = cur.fetchone()[0]
    return result if result is not None else 0
```

```python
res = cost_of('a')
print(res)
```

Output:
```
216.82
```

<!-- grade_id: load_5, points: 3, grade: true, solution: false -->
```python
res = cost_of('a')
assert type(res) == float
assert abs(res - 216.82) < 1e-5

try:
    cost_of(3)
except Exception:
    pass
else:
    raise Exception("should raise an exception")
```

## Problem 2: Building simple models with Pytorch (8 points) 
In this problem, you'll use **Pytorch** to build the loss functions for a pair of commonly-used statistical models. 

We will use variables $X$ and $Y$, which will serve as the predictor (independent variable) and response (dependent variable), respectively. Please use $W$ to denote a parameter that multiplies the predictor, and $b$ to denote a bias parameter (i.e., a parameter that is added).

**2(a)** (4 pts)

In this model, the binary variable $Y$ is distributed as a Bernoulli random variable with success parameter $\sigma(W^T X + b)$, where $\sigma(z) = (1+\exp(-z))^{-1}$ is the logistic function, $X \in R^6$ is the predictor random variable, and $W \in R^6, b \in R$ are the model parameters. 
 
Using **Pytorch** code, implement a class `LogisticRegression` that inherits from `nn.module`. This class should should have two attributes `w` and `b` which should be `nn.parameters` with shapes `(6,1)` and `(1)` respectivelly. 

This class should a method called `forward` that takes in the predictor random variable `x` with shape `(N, 6)`, where `N` is the number of observations, and returns the success parameter (also known as the prediction of our model on $Y$).

**Note:** Please initialize both `w, b` to be __all-one float tensors.__

<!-- grade_id: grade_21, points: 0, grade: false, solution: true -->
```python
import torch
import torch.nn as nn
class LogisticRegression(nn.Module):
    # YOUR CODE HERE
    def __init__(self):
        super(LogisticRegression, self).__init__()

        self.w = nn.Parameter(torch.randn(6,1))
        self.b = nn.Parameter(torch.randn(1))

    def forward(self, x):
        logits = x @ self.w + self.b
        probs = torch.sigmoid(logits)
        return probs
```

<!-- grade_id: load_21, points: 4, grade: true, solution: false -->
```python
model = LogisticRegression()
assert type(model.w) == nn.Parameter
assert type(model.b) == nn.Parameter
x = torch.eye(6).float()
y = model.forward(x)
assert y.shape == (6, 1)
```

**2(b)** (4 pts)

Using **Pytorch** code, write a function called `neg_log` that takes in `y_true` (true value of $Y$) and `y_pred` (predicted value of $Y$) and returns the negative log-likelihood loss function. You can assume that both `y_true` and `y_pred` have shapes `N x 1` where `N` is the number of observations. 
 
 __Hint:__ The loss should be a negative log-likelihood term, summed over all the observations. Remember that $Y$ is Bernoulli distributed which should suggest what the likelihood is. 

<!-- grade_id: grade_22, points: 0, grade: false, solution: true -->
```python
def neg_log(y_true, y_pred):
    # YOUR CODE HERE
    loss = -torch.sum(y_true * torch.log(y_pred) + (1 - y_true) * torch.log(1 - y_pred))
    return loss
```

<!-- grade_id: load_22, points: 4, grade: true, solution: false -->
```python
y_true = torch.tensor([1, 1, 1, 1, 1]).float().view(5, 1)
y_pred = torch.tensor([0.5, 0.5, 0.5, 0.5, 0.5]).float().view(5, 1)
assert 3.4 < neg_log(y_true, y_pred).item() < 3.6
```